---
title: "Hello, Tutorial!"
author: "Simon Knight, modified from J.J. Allaire"
date: "`r Sys.Date()`"
output: learnr::tutorial
runtime: shiny_prerendered
# Do not index/display tutorial by setting `private: true`
# private: true
description: >
  This is a demo tutorial.
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(learnrhash)
library(pacman)
p_load(kandinsky,morsecode, phonenumber, secrettext,tibble,ggplot2)

gradethis::gradethis_setup()
tutorial_options(exercise.checker = gradethis::grade_learnr)
##############
#converts letters to numbers
letter2num <- function(x) {utf8ToInt(x) - utf8ToInt("a") + 1L}
LETTER2num <- function(x) {utf8ToInt(x) - utf8ToInt("A") + 1L}
#unname(sapply("ABC", LETTER2num))
let2n <- function(x) { 
  
  x <- unlist(strsplit(x, split = ""))
  
  sapply(x, function(y) { 
    unlist(strsplit(y, split = ""))
  
  if(grepl("[[:upper:]]", y)) {
    LETTER2num(y)
} else {
    letter2num(y)
}
  }   ) }


tut_reptitle <- "Hello world"
```

The following code computes the answer to 1+1. Change it so it computes 2 + 2:

```{r addition, exercise=TRUE}
1 + 1
```

```{r addition-solution}
2 + 2
```

```{r addition-code-check}
grade_code()
```


